# Week 5 Exercises

_McKinney 3.2_

**Unlike in previous weeks, in the exercises below, you will need to create a function definition from scratch.  I'll provide specific instructions and examples for how it will be used, but you will have to do the work of definiging the whole function.**

---
---

### 25.1 Longest String

Write a function called `longest (L)` that takes as its only parameter a list of strings.  Your function needs to find the longest of those strings and return the position number of that longest string.  An example is provided below:

The longest word in that list of strings is "birthday" in position #1, so `longest(strings)` should return 1.
```
>>> strings = ['happy', 'birthday', 'to', 'me']
>>> longest(strings)
1
```

Be sure to include a docstring.  Include test cases in your docstring if you want to.

<!-- grade_id: longest-answer, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

def longest(strings):
    """(list) -> int
    This function finds the longest string in the list of input strings
    and returns the position number of that longest string.
    
    >>> longest(['happy', 'birthday', 'to', 'me'])
    1
    
    >>> longest(['one', 'two', 'three', 'four'])
    2
    """
    
    longest_pos = 0
    longest_len = 0
    
    print("Position  Item         Longest So Far")
    print("--------- ------------ --------------")
    for position, item in enumerate(strings):
        
        if len(item) > longest_len:
            longest_pos = position
            longest_len = len(item)
        print("{0:9} {1:12} {2} : {3}".format(position, item, longest_pos, longest_len))
            
    return longest_pos

### END SOLUTION
```

```python
longest(['happy', 'birthday', 'to', 'me'])
```

Output:
```
Position  Item         Longest So Far
--------- ------------ --------------
        0 happy        0 : 5
        1 birthday     1 : 8
        2 to           1 : 8
        3 me           1 : 8
1
```

<!-- grade_id: longest-tests, points: 10, grade: true, solution: false -->
```python
assert longest("happy birthday to me".split(" ")) == 1
assert longest("enjoy class".split(" ")) == 0
assert longest(['when','what','where','how','who']) == 2
```

Output:
```
Position  Item         Longest So Far
--------- ------------ --------------
        0 happy        0 : 5
        1 birthday     1 : 8
        2 to           1 : 8
        3 me           1 : 8
Position  Item         Longest So Far
--------- ------------ --------------
        0 enjoy        0 : 5
        1 class        0 : 5
Position  Item         Longest So Far
--------- ------------ --------------
        0 when         0 : 4
        1 what         0 : 4
        2 where        2 : 5
        3 how          2 : 5
        4 who          2 : 5
```

```python
import doctest
doctest.run_docstring_examples(longest, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    longest(['happy', 'birthday', 'to', 'me'])
Expecting:
    1
**********************************************************************
File "__main__", line 8, in NoName
Failed example:
    longest(['happy', 'birthday', 'to', 'me'])
Expected:
    1
Got:
    Position  Item         Longest So Far
    --------- ------------ --------------
            0 happy        0 : 5
            1 birthday     1 : 8
            2 to           1 : 8
            3 me           1 : 8
    1
Trying:
    longest(['one', 'two', 'three', 'four'])
Expecting:
    2
**********************************************************************
File "__main__", line 11, in NoName
Failed example:
    longest(['one', 'two', 'three', 'four'])
Expected:
    2
Got:
    Position  Item         Longest So Far
    --------- ------------ --------------
            0 one          0 : 3
            1 two          0 : 3
            2 three        2 : 5
            3 four         2 : 5
    2
```

### 25.2 Celsius to Farenheit

Write a function called f_to_c() that converts a given temperature in degrees Farenheit to degrees Celsius.  If you don't recall that conversion, it is:

$ temp_c = \frac{5}{9} \times (temp_f - 32) $

Make sure that you code is well documented using the DocString examples, and that your code includes tests for 212f, 32f, and 98.6f.

```python
### BEGIN SOLUTION
def f_to_c(degree_farenheit):
    
    """ (float(farenheit)) -> float(celsius) 
    This function takes an integer(temperature in farenheit) as the input and convert it to celsius.

    >>> f_to_c(212)
    100.0

    >>> f_to_c(32)
    0.0

    >>> f_to_c(98.6)
    37.0

    """
    ## raising and error if farenheit is less than 0
    
    if degree_farenheit < 0:
        raise ValueError('temperature should be a positive number. You gave me a {}'.format(print(degree_farenheit)))
        
    ## calculating celsius
    
    degree_celsius = 5/9*(degree_farenheit - 32)
    return degree_celsius
    

### END SOLUTION
```

```python
assert f_to_c(212) == 100.0
assert f_to_c(32) == 0.0
assert f_to_c(98.6) == 37.0
```

```python
import doctest
doctest.run_docstring_examples(f_to_c, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    f_to_c(212)
Expecting:
    100.0
ok
Trying:
    f_to_c(32)
Expecting:
    0.0
ok
Trying:
    f_to_c(98.6)
Expecting:
    37.0
ok
```

### 25.3 Computing Length of Stay

For this problem, we have a collection of patient enounter data stored as a Python dictionary.  The `key` for the dictionary is the **encounter ID**, a code that starts with the letter `E` followed by four numbers.  The value associated with each encounter ID is another Python dictionary.  This "inner" dictionary holds three items: admit date, primary diagnosis, and discharge date.  See the example in the code below.

You need to write a length of stay function that computes the length of stay, in whole days, between the admit date and discharge date.  However, if the diagnosis is "Observation" then the length of stay should always be returned as 0 regardless of the admit and discharge dates.

You will find it handy to refer to this example here on how to calculate the number of days between two dates: https://stackoverflow.com/questions/151199/how-to-calculate-number-of-days-between-two-given-dates

```python
from datetime import date

### BEGIN SOLUTION
def length_of_stay(admit, discharge, diagnosis):
    
    """(date, str, date) -> int
    This computes the length of stay of the patients, and if the diagnosis is just Observation, it computes the number of days as 0.
    days = discharge_date - admit_date 
    
    >>> length_of_stay(date(2021,1,9), date(2021,1,8), 'COPD')
    ValueError: admit > discharge
    
    >>> length_of_stay(date(2021,1,5), date(2021,1,9), 'Hypertension')
    4
    
    >>> length_of_stay(date(2021,1,23), date(2021,1,25), 'Observation')
    0
    """
    
    
    # for encounter_id, value in encounters.items():
    
    ## raising an error if admit_date is greater than discharge_date
    
    if admit > discharge:
        raise ValueError("admit > discharge ")
        
    ## calculating days of stay of patients 
    
        
    elif diagnosis != "Observation" and discharge >= admit:
        return (discharge - admit).days
    
    else:         
        return 0
    
### END SOLUTION
```

```python
import doctest
doctest.run_docstring_examples(length_of_stay, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    length_of_stay(date(2021,1,9), date(2021,1,8), 'COPD')
Expecting:
    ValueError: admit > discharge
**********************************************************************
File "__main__", line 10, in NoName
Failed example:
    length_of_stay(date(2021,1,9), date(2021,1,8), 'COPD')
Exception raised:
    Traceback (most recent call last):
      File "/opt/tljh/user/lib/python3.6/doctest.py", line 1330, in __run
        compileflags, 1), test.globs)
      File "<doctest NoName[0]>", line 1, in <module>
        length_of_stay(date(2021,1,9), date(2021,1,8), 'COPD')
      File "<ipython-input-8-6a5c76a69569>", line 26, in length_of_stay
        raise ValueError("admit > discharge ")
    ValueError: admit > discharge 
Trying:
    length_of_stay(date(2021,1,5), date(2021,1,9), 'Hypertension')
Expecting:
    4
ok
Trying:
    length_of_stay(date(2021,1,23), date(2021,1,25), 'Observation')
Expecting:
    0
ok
```

```python
encounters = { 
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}
for encounter_id, value in encounters.items():
    print(length_of_stay(value['admit'],value['discharge'],value['diagnosis']))


```

Output:
```
5
4
1
0
0
```

---
---

### 25.4 Average Length of Stay

Create a function called `average_los` that returns the average of the LOS for the encounters in the provided dictionary. The encounters must be a dictionary that contains dictionaries that each contain at least an admit date and a discharge date.
    
If the length of stay of any individual encounter is 0, then it will not be counted toward the average.

```python
test1 = {
        "E1234": { "admit": date(2019,1,10), "diagnosis": "COPD", "discharge": date(2019,1,8) }
}

test2 = {
        "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) }
}

test3 = {
        "E1234": { "admit": date(2019,1,5), "diagnosis": "COPD", "discharge": date(2019,1,8) }
}

 

### BEGIN SOLUTION
def average_los(dictionary):
    
    """(dictionary) -> float
    This computes the average days of stay of all the patients, taking atleast admit date and 
    discharge date as inputs in the dictionary.
    
    >>> average_los(test1)
    ValueError: admit > discharge
    
    >>> average_los(test2)
    5.0
    
    >>> average_los(test3)
    3.0
    
    """

   
    
    ## assigning variables to emoty list and 0.
    
    los = []
    avg_los = 0
    
    ## looping through the dictionary
    
    for encounter_id, value in dictionary.items():
        
    ## raising an error if admit_date is greater than discharge_date
        
        if value['admit'] > value['discharge']:
            raise ValueError("admit > discharge ")
            
    ## trying to check if discharge date is greater than admit date or not 
        
        try:
            value['discharge'] > value['admit']
            
    ## writing an if condition to calculate days of stay and appending all the items to the list
            
            if value['discharge'] > value['admit']:
                stay = (value['discharge'] - value['admit']).days
                los.append(stay)
            
    ## if the try condition fails i.e., if admit is equal to discharge then pass
    
        except:
            pass        
        
    ## calculating aveage length of stay of all the patients
        
    avg_los = sum(los)/len(los)
    return avg_los
    

### END SOLUTION
```

```python
from datetime import date
encounters = { 
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}

assert(average_los(encounters)) == 2.75
```

```python
print(average_los(encounters))
```

Output:
```
2.75
```

```python
import doctest
doctest.run_docstring_examples(average_los, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    average_los(test1)
Expecting:
    ValueError: admit > discharge
**********************************************************************
File "__main__", line 22, in NoName
Failed example:
    average_los(test1)
Exception raised:
    Traceback (most recent call last):
      File "/opt/tljh/user/lib/python3.6/doctest.py", line 1330, in __run
        compileflags, 1), test.globs)
      File "<doctest NoName[0]>", line 1, in <module>
        average_los(test1)
      File "<ipython-input-11-9c7df870e80b>", line 47, in average_los
        raise ValueError("admit > discharge ")
    ValueError: admit > discharge 
Trying:
    average_los(test2)
Expecting:
    5.0
ok
Trying:
    average_los(test3)
Expecting:
    3.0
ok
```

---

### 25.5 Celsius to Farenheit

Write your own function called `c_to_f` that converts degrees Celsius to degrees Farenheit.  Include in your solution a series of doc tests that can verify the conversion using inputs of 100, 0, and 37 degrees Celsius

```python
### BEGIN SOLUTION
def c_to_f(degree_celsius):
    """(int(celsius)) -> (float(farenheit))
    This function coverts given inputs in degree celsius to farenheit.
    
    >>> c_to_f(100)
    212.0
    
    >>> c_to_f(0)
    32.0
    
    >>> c_to_f(37)
    98.60000000000001
    
    """
    ## calculating farenheit degrees, if the celcius is given
    
    degree_farenheit = (9/5*degree_celsius) + 32 
    return degree_farenheit
### END SOLUTION
```

```python
import doctest
doctest.run_docstring_examples(c_to_f, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    c_to_f(100)
Expecting:
    212.0
ok
Trying:
    c_to_f(0)
Expecting:
    32.0
ok
Trying:
    c_to_f(37)
Expecting:
    98.60000000000001
ok
```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to use the `git` command line program to **add** your homework file (this file) to your local repository, **commit** your changes to your local repository, and then **push** those changes up to github.com.  From there, I'll be able to **pull** the changes down and do my grading.  I'll provide some feedback, **commit** and **push** my comments back to you.  Next week, I'll show you how to **pull** down my comments.

To run through everything one last time and submit your work:
1. Use the `Kernel` -> `Restart Kernel and Run All Cells` menu option to run everything from top to bottom and stop here.
2. Follow the instruction on the prompt below to either ssave and submit your work, or continue working.

If anything fails along the way with this submission part of the process, let me know.  I'll help you troubleshoort.

---

```python
a=input('''
Are you ready to submit your work?
1. Click the Save icon (or do Ctrl-S / Cmd-S)
2. Type "yes" or "no" below
3. Press Enter

''')

if a=='yes':
    !git add week05_assignment_2.ipynb
    !git commit -a -m "Submitting the week 5 programming exercises"
    !git push
else:
    print('''
    
OK. We can wait.
''')
```

Output:
```
Are you ready to submit your work?
1. Click the Save icon (or do Ctrl-S / Cmd-S)
2. Type "yes" or "no" below
3. Press Enter

 yes
[main 4aa03fa] Submitting the week 5 programming exercises
 2 files changed, 842 insertions(+), 2 deletions(-)
 create mode 100644 week05/week05_assignment_2.ipynb
Warning: Permanently added the ECDSA host key for IP address '140.82.112.3' to the list of known hosts.
Counting objects: 5, done.
Delta compression using up to 2 threads.
Compressing objects: 100% (5/5), done.
Writing objects: 100% (5/5), 5.90 KiB | 5.90 MiB/s, done.
Total 5 (delta 2), reused 0 (delta 0)
remote: Resolving deltas: 100% (2/2), completed with 2 local objects.[K
remote: This repository moved. Please use the new location:[K
remote:   git@github.com:Sravani9544/hds5210-2022.git[K
To github.com:sravani9544/hds5210-2022.git
   46f21ac..4aa03fa  main -> main
```


---

If the message above says something like _Submitting the week 3 review exercises_ or _Everything is up to date_, then your work was submitted correctly.

```python

```